(assume a0 (not (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))))
(step t2 (cl (= (not (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (and (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$))) (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))))))) :rule ite_intro)
(step t3 (cl (not (= (not (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (and (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$))) (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))))))) (not (not (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))))) (and (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$))) (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))))) :rule equiv_pos2)
(step t4 (cl (not (not (not (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))))) (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) :rule not_not)
(step t5 (cl (not (= (not (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (and (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$))) (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))))))) (= (fun_app$ (ite c$ f$ g$) x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (and (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$))) (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))))) :rule th_resolution :premises (t4 t3))
(step t6 (cl (and (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$))) (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))))) :rule th_resolution :premises (a0 t2 t5))
(step t7 (cl (not (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$)))) :rule and :premises (t6))
(step t8 (cl (ite c$ (= f$ (ite c$ f$ g$)) (= g$ (ite c$ f$ g$)))) :rule and :premises (t6))
(step t9 (cl c$ (= g$ (ite c$ f$ g$))) :rule ite1 :premises (t8))
(step t10 (cl (not c$) (= f$ (ite c$ f$ g$))) :rule ite2 :premises (t8))
(step t11 (cl (ite c$ (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))) (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$))))) :rule and :premises (t6))
(step t12 (cl c$ (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) :rule ite1 :premises (t11))
(step t13 (cl (not c$) (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) :rule ite2 :premises (t11))
(step t14 (cl (not (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (not (= (fun_app$ g$ x$) (fun_app$ (ite c$ f$ g$) x$))) (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) :rule eq_transitive)
(step t15 (cl (not (= g$ (ite c$ f$ g$))) (not (= x$ x$)) (= (fun_app$ g$ x$) (fun_app$ (ite c$ f$ g$) x$))) :rule eq_congruent)
(step t16 (cl (= x$ x$)) :rule eq_reflexive)
(step t17 (cl (not (= g$ (ite c$ f$ g$))) (= (fun_app$ g$ x$) (fun_app$ (ite c$ f$ g$) x$))) :rule th_resolution :premises (t15 t16))
(step t18 (cl (not (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$)) (not (= g$ (ite c$ f$ g$)))) :rule th_resolution :premises (t14 t17))
(step t19 (cl (not (= (fun_app$ g$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (not (= g$ (ite c$ f$ g$)))) :rule resolution :premises (t18 t7))
(step t20 (cl c$) :rule resolution :premises (t19 t12 t9))
(step t21 (cl (= f$ (ite c$ f$ g$))) :rule resolution :premises (t10 t20))
(step t22 (cl (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) :rule resolution :premises (t13 t20))
(step t23 (cl (not (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (not (= (fun_app$ f$ x$) (fun_app$ (ite c$ f$ g$) x$))) (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$))) :rule eq_transitive)
(step t24 (cl (not (= f$ (ite c$ f$ g$))) (not (= x$ x$)) (= (fun_app$ f$ x$) (fun_app$ (ite c$ f$ g$) x$))) :rule eq_congruent)
(step t25 (cl (not (= f$ (ite c$ f$ g$))) (= (fun_app$ f$ x$) (fun_app$ (ite c$ f$ g$) x$))) :rule th_resolution :premises (t24 t16))
(step t26 (cl (not (= (fun_app$ f$ x$) (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)))) (= (ite c$ (fun_app$ f$ x$) (fun_app$ g$ x$)) (fun_app$ (ite c$ f$ g$) x$)) (not (= f$ (ite c$ f$ g$)))) :rule th_resolution :premises (t23 t25))
(step t27 (cl) :rule resolution :premises (t26 t7 t21 t22))
