(assume a0 (or a c))
(assume a1 (or a (not c) d))
(assume a2 (or (not d) e))
(assume a3 (or (not d) (not e)))
(assume a4 (not a))
(assume a5 (not b))
(step t0 (cl a c) :rule or :premises (a0))
(step t1 (cl a (not c) d) :rule or :premises (a1))
(step t2 (cl (not d) e) :rule or :premises (a2))
(step t3 (cl (not d) (not e)) :rule or :premises (a3))
(step t4 (cl (not d)) :rule resolution :premises (t2 t3))
(step t4.t2 (cl a (not c)) :rule resolution :premises (t1 t4))
(step t4.t2.t3 (cl a b) :rule resolution :premises (t0 t4.t2))
(step t10 (cl) :rule resolution :premises (t4.t2.t3 a4 a5) :args (a true b true))