(assume a0 (= a b))
(assume a1 (= c d))
(assume a4 (or (not p3) (not (= (f a c) (f b d)))))
(step t7 (cl (not (= a b)) (not (= c d)) (= (f a c) (f b d))) :rule eq_congruent)
(step t8 (cl (= (f a c) (f b d)) (not (= a b)) (not (= c d))) :rule reordering :premises (t7))
(step t10 (cl (not p3) (not (= (f a c) (f b d)))) :rule or :premises (a4))
(step t11.t2 (cl (= (f a c) (f b d)) (not (= c d))) :rule resolution :premises (t8 a0))
(step t11.t2.t3 (cl (= (f a c) (f b d))) :rule resolution :premises (a1 t11.t2))
(step t11.t2.t3.t4.t5.t6.t7.t8.t9.t9 (cl (not p3)) :rule resolution :premises (t10 t11.t2.t3))
(step t11.t2.t3.t4.t5.t6.t7.t8.t9.t9.t10 (cl) :rule resolution :premises (t11.t2.t3.t4.t5.t6.t7.t8.t9.t9))