(assume a0 (not (= (times$ (f$ (fun_app$ g$ x$)) (vector_derivative$ g$ (at_within$ x$ top$))) (times$ (f$a (fun_app$ g$a x$)) (vector_derivative$ g$a (at_within$ x$ top$))))))
(assume a1 (= g$ g$a))
(assume a2 (= (f$a (fun_app$ g$a x$)) (f$ (fun_app$ g$a x$))))
(step t4 (cl (not (= (f$ (fun_app$ g$ x$)) (f$a (fun_app$ g$a x$)))) (not (= (vector_derivative$ g$ (at_within$ x$ top$)) (vector_derivative$ g$a (at_within$ x$ top$)))) (= (times$ (f$ (fun_app$ g$ x$)) (vector_derivative$ g$ (at_within$ x$ top$))) (times$ (f$a (fun_app$ g$a x$)) (vector_derivative$ g$a (at_within$ x$ top$))))) :rule eq_congruent)
(step t5 (cl (not (= (f$ (fun_app$ g$ x$)) (f$ (fun_app$ g$a x$)))) (not (= (f$a (fun_app$ g$a x$)) (f$ (fun_app$ g$a x$)))) (= (f$ (fun_app$ g$ x$)) (f$a (fun_app$ g$a x$)))) :rule eq_transitive)
(step t6 (cl (not (= (fun_app$ g$ x$) (fun_app$ g$a x$))) (= (f$ (fun_app$ g$ x$)) (f$ (fun_app$ g$a x$)))) :rule eq_congruent)
(step t7 (cl (not (= g$ g$a)) (not (= x$ x$)) (= (fun_app$ g$ x$) (fun_app$ g$a x$))) :rule eq_congruent)
(step t8 (cl (= x$ x$)) :rule eq_reflexive)
(step t9 (cl (not (= g$ g$a)) (= (fun_app$ g$ x$) (fun_app$ g$a x$))) :rule th_resolution :premises (t7 t8))
(step t10 (cl (= (f$ (fun_app$ g$ x$)) (f$ (fun_app$ g$a x$))) (not (= g$ g$a))) :rule th_resolution :premises (t6 t9))
(step t11 (cl (not (= (f$a (fun_app$ g$a x$)) (f$ (fun_app$ g$a x$)))) (= (f$ (fun_app$ g$ x$)) (f$a (fun_app$ g$a x$))) (not (= g$ g$a))) :rule th_resolution :premises (t5 t10))
(step t12 (cl (not (= (vector_derivative$ g$ (at_within$ x$ top$)) (vector_derivative$ g$a (at_within$ x$ top$)))) (= (times$ (f$ (fun_app$ g$ x$)) (vector_derivative$ g$ (at_within$ x$ top$))) (times$ (f$a (fun_app$ g$a x$)) (vector_derivative$ g$a (at_within$ x$ top$)))) (not (= (f$a (fun_app$ g$a x$)) (f$ (fun_app$ g$a x$)))) (not (= g$ g$a))) :rule th_resolution :premises (t4 t11))
(step t13 (cl (not (= g$ g$a)) (not (= (at_within$ x$ top$) (at_within$ x$ top$))) (= (vector_derivative$ g$ (at_within$ x$ top$)) (vector_derivative$ g$a (at_within$ x$ top$)))) :rule eq_congruent)
(step t14 (cl (= (at_within$ x$ top$) (at_within$ x$ top$))) :rule eq_reflexive)
(step t15 (cl (not (= g$ g$a)) (= (vector_derivative$ g$ (at_within$ x$ top$)) (vector_derivative$ g$a (at_within$ x$ top$)))) :rule th_resolution :premises (t13 t14))
(step t16 (cl (= (times$ (f$ (fun_app$ g$ x$)) (vector_derivative$ g$ (at_within$ x$ top$))) (times$ (f$a (fun_app$ g$a x$)) (vector_derivative$ g$a (at_within$ x$ top$)))) (not (= (f$a (fun_app$ g$a x$)) (f$ (fun_app$ g$a x$)))) (not (= g$ g$a)) (not (= g$ g$a))) :rule th_resolution :premises (t12 t15))
(step t17 (cl (= (times$ (f$ (fun_app$ g$ x$)) (vector_derivative$ g$ (at_within$ x$ top$))) (times$ (f$a (fun_app$ g$a x$)) (vector_derivative$ g$a (at_within$ x$ top$)))) (not (= (f$a (fun_app$ g$a x$)) (f$ (fun_app$ g$a x$)))) (not (= g$ g$a))) :rule contraction :premises (t16))
(step t18 (cl) :rule resolution :premises (t17 a0 a1 a2))
