(assume a0 (not (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ (plus$ d0$ (image_mset$ neg$ (mset$ as0$)))))))
(assume a1 (= (plus$ d0$ (image_mset$ neg$ (mset$ as0$))) dA0$))
(assume a2 (=> (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))))
(assume a3 (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)))
(step t5 (cl (not (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$))) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) :rule implies :premises (a2))
(step t6 (cl (not (not (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)))) (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)) :rule not_not)
(step t7 (cl (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) :rule th_resolution :premises (t6 t5))
(step t8 (cl (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) :rule resolution :premises (t7 a3))
(step t9 (cl (not (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) (not (= (fun_app$ s$ (plus$ d0$ (image_mset$ neg$ (mset$ as0$)))) (fun_app$ s$ dA0$))) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ (plus$ d0$ (image_mset$ neg$ (mset$ as0$)))))) :rule eq_transitive)
(step t10 (cl (not (= s$ s$)) (not (= (plus$ d0$ (image_mset$ neg$ (mset$ as0$))) dA0$)) (= (fun_app$ s$ (plus$ d0$ (image_mset$ neg$ (mset$ as0$)))) (fun_app$ s$ dA0$))) :rule eq_congruent)
(step t11 (cl (= s$ s$)) :rule eq_reflexive)
(step t12 (cl (not (= (plus$ d0$ (image_mset$ neg$ (mset$ as0$))) dA0$)) (= (fun_app$ s$ (plus$ d0$ (image_mset$ neg$ (mset$ as0$)))) (fun_app$ s$ dA0$))) :rule th_resolution :premises (t10 t11))
(step t13 (cl (not (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ (plus$ d0$ (image_mset$ neg$ (mset$ as0$))))) (not (= (plus$ d0$ (image_mset$ neg$ (mset$ as0$))) dA0$))) :rule th_resolution :premises (t9 t12))
(step t14 (cl) :rule resolution :premises (t13 a0 a1 t8))
