(assume a0 (not (= (image_mset$ neg$ (mset$ as$)) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) eta_0$))))
(assume a1 (= (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) eta_0$) (s_M$ subst_atm$ s$ m$ da$)))
(assume a2 (= da$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))
(assume a3 (= (image_mset$ neg$ (mset$ as$)) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$))))))
(step t5 (cl (not (= (image_mset$ neg$ (mset$ as$)) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) (not (= (s_M$ subst_atm$ s$ m$ da$) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) (not (= (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) eta_0$) (s_M$ subst_atm$ s$ m$ da$))) (= (image_mset$ neg$ (mset$ as$)) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) eta_0$))) :rule eq_transitive)
(step t6 (cl (not (= subst_atm$ subst_atm$)) (not (= s$ s$)) (not (= m$ m$)) (not (= da$ (plus$ d$ (image_mset$ neg$ (mset$ as$))))) (= (s_M$ subst_atm$ s$ m$ da$) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) :rule eq_congruent)
(step t7 (cl (= subst_atm$ subst_atm$)) :rule eq_reflexive)
(step t8 (cl (not (= s$ s$)) (not (= m$ m$)) (not (= da$ (plus$ d$ (image_mset$ neg$ (mset$ as$))))) (= (s_M$ subst_atm$ s$ m$ da$) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) :rule th_resolution :premises (t6 t7))
(step t9 (cl (= s$ s$)) :rule eq_reflexive)
(step t10 (cl (not (= m$ m$)) (not (= da$ (plus$ d$ (image_mset$ neg$ (mset$ as$))))) (= (s_M$ subst_atm$ s$ m$ da$) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) :rule th_resolution :premises (t8 t9))
(step t11 (cl (= m$ m$)) :rule eq_reflexive)
(step t12 (cl (not (= da$ (plus$ d$ (image_mset$ neg$ (mset$ as$))))) (= (s_M$ subst_atm$ s$ m$ da$) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) :rule th_resolution :premises (t10 t11))
(step t13 (cl (not (= (image_mset$ neg$ (mset$ as$)) (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) (not (= (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) eta_0$) (s_M$ subst_atm$ s$ m$ da$))) (= (image_mset$ neg$ (mset$ as$)) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) eta_0$)) (not (= da$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))))) :rule th_resolution :premises (t5 t12))
(step t14 (cl) :rule resolution :premises (t13 a0 a1 a2 a3))
