(assume a0 (not (= (subst_cls$ subst_atm$ (image_mset$ neg$ (mset$ as0$)) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))))))
(assume a1 (=> (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))))
(assume a2 (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)))
(step t4 (cl (not (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$))) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) :rule implies :premises (a1))
(step t5 (cl (not (not (not (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)))) (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$)) :rule not_not)
(step t6 (cl (= (s_M$ subst_atm$ s$ m$ (plus$ d$ (image_mset$ neg$ (mset$ as$)))) zero$) (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) :rule th_resolution :premises (t5 t4))
(step t7 (cl (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) :rule resolution :premises (t6 a2))
(step t8 (cl (not (= subst_atm$ subst_atm$)) (not (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) (not (= (hd$ (renamings_apart$ (cons$ dA0$ cAs0$))) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$))))) (= (subst_cls$ subst_atm$ (image_mset$ neg$ (mset$ as0$)) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))))) :rule eq_congruent)
(step t9 (cl (= subst_atm$ subst_atm$)) :rule eq_reflexive)
(step t10 (cl (not (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) (not (= (hd$ (renamings_apart$ (cons$ dA0$ cAs0$))) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$))))) (= (subst_cls$ subst_atm$ (image_mset$ neg$ (mset$ as0$)) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))))) :rule th_resolution :premises (t8 t9))
(step t11 (cl (= (hd$ (renamings_apart$ (cons$ dA0$ cAs0$))) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$))))) :rule eq_reflexive)
(step t12 (cl (not (= (image_mset$ neg$ (mset$ as0$)) (fun_app$ s$ dA0$))) (= (subst_cls$ subst_atm$ (image_mset$ neg$ (mset$ as0$)) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))) (subst_cls$ subst_atm$ (fun_app$ s$ dA0$) (hd$ (renamings_apart$ (cons$ dA0$ cAs0$)))))) :rule th_resolution :premises (t10 t11))
(step t13 (cl) :rule resolution :premises (t12 a0 t7))
