(assume a0 (not (not (and (subsumes$ subst_atm$ d$ c$) (not (subsumes$ subst_atm$ c$ d$))))))
(step t2 (cl (not (not (not (and (subsumes$ subst_atm$ d$ c$) (not (subsumes$ subst_atm$ c$ d$)))))) (and (subsumes$ subst_atm$ d$ c$) (not (subsumes$ subst_atm$ c$ d$)))) :rule not_not)
(step t3 (cl (and (subsumes$ subst_atm$ d$ c$) (not (subsumes$ subst_atm$ c$ d$)))) :rule th_resolution :premises (t2 a0))
(assume a1 (and (subsumes$ subst_atm$ c$ d$) (not (subsumes$ subst_atm$ d$ c$))))
(step t5 (cl (not (subsumes$ subst_atm$ c$ d$))) :rule and :premises (t3))
(step t6 (cl (subsumes$ subst_atm$ c$ d$)) :rule and :premises (a1))
(step t7 (cl) :rule resolution :premises (t6 t5))
