(assume a1 (not (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))))))
(assume a2 (chain$ (derive$ ground_sound_Gamma$ rf$) gs$))
(assume a3 (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$)))))
(assume a4 (=> (chain$ (derive$ ground_sound_Gamma$ rf$) gs$) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))))
(assume a5 (= (lmap$ uu$ sts$) gs$))
(step t6 (cl (not (chain$ (derive$ ground_sound_Gamma$ rf$) gs$)) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) :rule implies :premises (a4))
(step t7 (cl (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) :rule resolution :premises (t6 a2))
(step t8 (cl (not (= gamma$ gamma$)) (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))))) (not (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))))) (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))))) :rule eq_congruent_pred)
(step t9 (cl (= gamma$ gamma$)) :rule eq_reflexive)
(step t10 (cl (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))))) (not (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))))) (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))))) :rule th_resolution :premises (t8 t9))
(step t11 (cl (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)))) (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))))) :rule eq_transitive)
(step t12 (cl (not (= s$ s$)) (not (= subst_atm$ subst_atm$)) (not (= sts$ sts$)) (not (= (liminf_llist$ (lmap$ uu$ sts$)) (liminf_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)))) :rule eq_congruent)
(step t13 (cl (= s$ s$)) :rule eq_reflexive)
(step t14 (cl (not (= subst_atm$ subst_atm$)) (not (= sts$ sts$)) (not (= (liminf_llist$ (lmap$ uu$ sts$)) (liminf_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)))) :rule th_resolution :premises (t12 t13))
(step t15 (cl (= subst_atm$ subst_atm$)) :rule eq_reflexive)
(step t16 (cl (not (= sts$ sts$)) (not (= (liminf_llist$ (lmap$ uu$ sts$)) (liminf_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)))) :rule th_resolution :premises (t14 t15))
(step t17 (cl (= sts$ sts$)) :rule eq_reflexive)
(step t18 (cl (not (= (liminf_llist$ (lmap$ uu$ sts$)) (liminf_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)))) :rule th_resolution :premises (t16 t17))
(step t19 (cl (not (= (lmap$ uu$ sts$) gs$)) (= (liminf_llist$ (lmap$ uu$ sts$)) (liminf_llist$ gs$))) :rule eq_congruent)
(step t20 (cl (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$))) (not (= (lmap$ uu$ sts$) gs$))) :rule th_resolution :premises (t18 t19))
(step t21 (cl (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$)))) (not (= (lmap$ uu$ sts$) gs$))) :rule th_resolution :premises (t11 t20))
(step t22 (cl (not (= s$ s$)) (not (= subst_atm$ subst_atm$)) (not (= sts$ sts$)) (not (= (sup_llist$ (lmap$ uu$ sts$)) (sup_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) :rule eq_congruent)
(step t23 (cl (not (= subst_atm$ subst_atm$)) (not (= sts$ sts$)) (not (= (sup_llist$ (lmap$ uu$ sts$)) (sup_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) :rule th_resolution :premises (t22 t13))
(step t24 (cl (not (= sts$ sts$)) (not (= (sup_llist$ (lmap$ uu$ sts$)) (sup_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) :rule th_resolution :premises (t23 t15))
(step t25 (cl (not (= (sup_llist$ (lmap$ uu$ sts$)) (sup_llist$ gs$))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) :rule th_resolution :premises (t24 t17))
(step t26 (cl (not (= (lmap$ uu$ sts$) gs$)) (= (sup_llist$ (lmap$ uu$ sts$)) (sup_llist$ gs$))) :rule eq_congruent)
(step t27 (cl (= (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$))) (not (= (lmap$ uu$ sts$) gs$))) :rule th_resolution :premises (t25 t26))
(step t28 (cl (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$)))) (not (= (lmap$ uu$ sts$) gs$)) (not (= (lmap$ uu$ sts$) gs$))) :rule th_resolution :premises (t21 t27))
(step t29 (cl (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$))) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$)))) (not (= (lmap$ uu$ sts$) gs$))) :rule contraction :premises (t28))
(step t30 (cl (not (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ (lmap$ uu$ sts$))))) (member$ gamma$ (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ (lmap$ uu$ sts$)))) (not (= (sr_ext_Ri$ s$ subst_atm$ sts$ (liminf_llist$ gs$)) (sr_ext_Ri$ s$ subst_atm$ sts$ (sup_llist$ gs$)))) (not (= (lmap$ uu$ sts$) gs$))) :rule th_resolution :premises (t10 t29))
(step t31 (cl) :rule resolution :premises (t30 a1 a3 t7 a5))
