(assume a0 (not (not (= da$ zero$))))
(step t2 (cl (not (not (not (= da$ zero$)))) (= da$ zero$)) :rule not_not)
(step t3 (cl (= da$ zero$)) :rule th_resolution :premises (t2 a0))
(assume a1 (member$ da$ n$))
(assume a2 (not (member$ zero$ n$)))
(step t6 (cl (not (= da$ zero$)) (not (= n$ n$)) (not (member$ da$ n$)) (member$ zero$ n$)) :rule eq_congruent_pred)
(step t7 (cl (= n$ n$)) :rule eq_reflexive)
(step t8 (cl (not (= da$ zero$)) (not (member$ da$ n$)) (member$ zero$ n$)) :rule th_resolution :premises (t6 t7))
(step t9 (cl) :rule resolution :premises (t8 t3 a1 a2))
