(assume h1 (p a))
(assume h2 (not (= b (ite (p a) b a))))
(step t3 (cl (= (not (= b (ite (p a) b a))) (and (not (= b (ite (p a) b a))) (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a)))))) :rule ite_intro)
(step t4 (cl (not (= (not (= b (ite (p a) b a))) (and (not (= b (ite (p a) b a))) (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a)))))) (not (not (= b (ite (p a) b a)))) (and (not (= b (ite (p a) b a))) (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a))))) :rule equiv_pos2)
(step t5 (cl (not (not (not (= b (ite (p a) b a))))) (= b (ite (p a) b a))) :rule not_not)
(step t6 (cl (not (= (not (= b (ite (p a) b a))) (and (not (= b (ite (p a) b a))) (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a)))))) (= b (ite (p a) b a)) (and (not (= b (ite (p a) b a))) (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a))))) :rule th_resolution :premises (t5 t4))
(step t7 (cl (and (not (= b (ite (p a) b a))) (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a))))) :rule th_resolution :premises (h2 t3 t6))
(step t8 (cl (not (= b (ite (p a) b a)))) :rule and :premises (t7))
(step t9 (cl (ite (p a) (= b (ite (p a) b a)) (= a (ite (p a) b a)))) :rule and :premises (t7))
(step t10 (cl (p a) (= a (ite (p a) b a))) :rule ite1 :premises (t9))
(step t11 (cl (not (p a)) (= b (ite (p a) b a))) :rule ite2 :premises (t9))
(step t12 (cl) :rule resolution :premises (t11 h1 t8))
