(assume a3 (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$)))))
(assume a4 (= (plus$ i$ k2$) zero$))
(assume a5 (not (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ zero$)))))
(step t4 (cl (not (= (size$ u$) (size$ u$))) (not (= (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))) (size$ (worst_chain$ uu$ uua$ zero$)))) (not (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))))) (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ zero$)))) :rule eq_congruent_pred)
(step t5 (cl (= (size$ u$) (size$ u$))) :rule eq_reflexive)
(step t6 (cl (not (= (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))) (size$ (worst_chain$ uu$ uua$ zero$)))) (not (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))))) (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ zero$)))) :rule th_resolution :premises (t4 t5))
(step t7 (cl (not (= (worst_chain$ uu$ uua$ (plus$ i$ k2$)) (worst_chain$ uu$ uua$ zero$))) (= (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))) (size$ (worst_chain$ uu$ uua$ zero$)))) :rule eq_congruent)
(step t8 (cl (not (= uu$ uu$)) (not (= uua$ uua$)) (not (= (plus$ i$ k2$) zero$)) (= (worst_chain$ uu$ uua$ (plus$ i$ k2$)) (worst_chain$ uu$ uua$ zero$))) :rule eq_congruent)
(step t9 (cl (= uu$ uu$)) :rule eq_reflexive)
(step t10 (cl (not (= uua$ uua$)) (not (= (plus$ i$ k2$) zero$)) (= (worst_chain$ uu$ uua$ (plus$ i$ k2$)) (worst_chain$ uu$ uua$ zero$))) :rule th_resolution :premises (t8 t9))
(step t11 (cl (= uua$ uua$)) :rule eq_reflexive)
(step t12 (cl (not (= (plus$ i$ k2$) zero$)) (= (worst_chain$ uu$ uua$ (plus$ i$ k2$)) (worst_chain$ uu$ uua$ zero$))) :rule th_resolution :premises (t10 t11))
(step t13 (cl (= (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))) (size$ (worst_chain$ uu$ uua$ zero$))) (not (= (plus$ i$ k2$) zero$))) :rule th_resolution :premises (t7 t12))
(step t14 (cl (not (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ (plus$ i$ k2$))))) (less$ (size$ u$) (size$ (worst_chain$ uu$ uua$ zero$))) (not (= (plus$ i$ k2$) zero$))) :rule th_resolution :premises (t6 t13))
(step t15 (cl) :rule resolution :premises (t14 a3 a4 a5))
