(assume a0 (not (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$)))))
(step t2 (cl (not (not (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))))) (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) :rule not_not)
(step t3 (cl (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) :rule th_resolution :premises (t2 a0))
(assume a1 (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))))
(assume a2 (= (head$ sa$) (head$ ta$)))
(step t6 (cl (not (= gt_sym$ gt_sym$)) (not (= ground_heads_var$ ground_heads_var$)) (not (= (head$ ua$) (head$ ua$))) (not (= (head$ sa$) (head$ ta$))) (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) :rule eq_congruent_pred)
(step t7 (cl (= gt_sym$ gt_sym$)) :rule eq_reflexive)
(step t8 (cl (not (= ground_heads_var$ ground_heads_var$)) (not (= (head$ ua$) (head$ ua$))) (not (= (head$ sa$) (head$ ta$))) (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) :rule th_resolution :premises (t6 t7))
(step t9 (cl (= ground_heads_var$ ground_heads_var$)) :rule eq_reflexive)
(step t10 (cl (not (= (head$ ua$) (head$ ua$))) (not (= (head$ sa$) (head$ ta$))) (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) :rule th_resolution :premises (t8 t9))
(step t11 (cl (= (head$ ua$) (head$ ua$))) :rule eq_reflexive)
(step t12 (cl (not (= (head$ sa$) (head$ ta$))) (not (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) (comp_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) :rule th_resolution :premises (t10 t11))
(step t13 (cl) :rule resolution :premises (t12 t3 a1 a2))
