(assume a0 (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))))
(assume a1 (= sa$ t2$))
(assume a2 (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$)))
(step t4 (cl (not (= ground_heads_var$ ground_heads_var$)) (not (= gt_sym$ gt_sym$)) (not (= delta$ delta$)) (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule eq_congruent_pred)
(step t5 (cl (= ground_heads_var$ ground_heads_var$)) :rule eq_reflexive)
(step t6 (cl (not (= gt_sym$ gt_sym$)) (not (= delta$ delta$)) (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t4 t5))
(step t7 (cl (= gt_sym$ gt_sym$)) :rule eq_reflexive)
(step t8 (cl (not (= delta$ delta$)) (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t6 t7))
(step t9 (cl (= delta$ delta$)) :rule eq_reflexive)
(step t10 (cl (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t8 t9))
(step t11 (cl (= extf$ extf$)) :rule eq_reflexive)
(step t12 (cl (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t10 t11))
(step t13 (cl (= arity_sym$ arity_sym$)) :rule eq_reflexive)
(step t14 (cl (not (= wt_sym$ wt_sym$)) (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t12 t13))
(step t15 (cl (= wt_sym$ wt_sym$)) :rule eq_reflexive)
(step t16 (cl (not (= (subst$ rho$ ta$) (subst$ rho$ ta$))) (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t14 t15))
(step t17 (cl (= (subst$ rho$ ta$) (subst$ rho$ ta$))) :rule eq_reflexive)
(step t18 (cl (not (= (subst$ rho$ sa$) (subst$ rho$ t2$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$))) :rule th_resolution :premises (t16 t17))
(step t19 (cl (not (= rho$ rho$)) (not (= sa$ t2$)) (= (subst$ rho$ sa$) (subst$ rho$ t2$))) :rule eq_congruent)
(step t20 (cl (= rho$ rho$)) :rule eq_reflexive)
(step t21 (cl (not (= sa$ t2$)) (= (subst$ rho$ sa$) (subst$ rho$ t2$))) :rule th_resolution :premises (t19 t20))
(step t22 (cl (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ t2$))) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (subst$ rho$ ta$) (subst$ rho$ sa$)) (not (= sa$ t2$))) :rule th_resolution :premises (t18 t21))
(step t23 (cl) :rule resolution :premises (t22 a0 a1 a2))
