(assume a0 (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))))
(assume a1 (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$)))
(assume a2 (= (head$ sa$) (head$ ta$)))
(step t4 (cl (not (= gt_sym$ gt_sym$)) (not (= ground_heads_var$ ground_heads_var$)) (not (= (head$ ua$) (head$ ua$))) (not (= (head$ sa$) (head$ ta$))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) :rule eq_congruent_pred)
(step t5 (cl (= gt_sym$ gt_sym$)) :rule eq_reflexive)
(step t6 (cl (not (= ground_heads_var$ ground_heads_var$)) (not (= (head$ ua$) (head$ ua$))) (not (= (head$ sa$) (head$ ta$))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) :rule th_resolution :premises (t4 t5))
(step t7 (cl (= ground_heads_var$ ground_heads_var$)) :rule eq_reflexive)
(step t8 (cl (not (= (head$ ua$) (head$ ua$))) (not (= (head$ sa$) (head$ ta$))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) :rule th_resolution :premises (t6 t7))
(step t9 (cl (= (head$ ua$) (head$ ua$))) :rule eq_reflexive)
(step t10 (cl (not (= (head$ sa$) (head$ ta$))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ ta$))) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ua$) (head$ sa$))) :rule th_resolution :premises (t8 t9))
(step t11 (cl) :rule resolution :premises (t10 a0 a1 a2))
