(assume a0 (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)))
(assume a1 (= t2$ (arg$ ta$)))
(assume a2 (or (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$) (= sa$ (arg$ ta$))))
(assume a3 (not (= t2$ sa$)))
(step t5 (cl (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$) (= sa$ (arg$ ta$))) :rule or :premises (a2))
(step t6 (cl (not (= ground_heads_var$ ground_heads_var$)) (not (= gt_sym$ gt_sym$)) (not (= delta$ delta$)) (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule eq_congruent_pred)
(step t7 (cl (= ground_heads_var$ ground_heads_var$)) :rule eq_reflexive)
(step t8 (cl (not (= gt_sym$ gt_sym$)) (not (= delta$ delta$)) (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t6 t7))
(step t9 (cl (= gt_sym$ gt_sym$)) :rule eq_reflexive)
(step t10 (cl (not (= delta$ delta$)) (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t8 t9))
(step t11 (cl (= delta$ delta$)) :rule eq_reflexive)
(step t12 (cl (not (= extf$ extf$)) (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t10 t11))
(step t13 (cl (= extf$ extf$)) :rule eq_reflexive)
(step t14 (cl (not (= arity_sym$ arity_sym$)) (not (= wt_sym$ wt_sym$)) (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t12 t13))
(step t15 (cl (= arity_sym$ arity_sym$)) :rule eq_reflexive)
(step t16 (cl (not (= wt_sym$ wt_sym$)) (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t14 t15))
(step t17 (cl (= wt_sym$ wt_sym$)) :rule eq_reflexive)
(step t18 (cl (not (= t2$ (arg$ ta$))) (not (= sa$ sa$)) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t16 t17))
(step t19 (cl (= sa$ sa$)) :rule eq_reflexive)
(step t20 (cl (not (= t2$ (arg$ ta$))) (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$)) (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ t2$ sa$)) :rule th_resolution :premises (t18 t19))
(step t21 (cl (not (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ (arg$ ta$) sa$))) :rule resolution :premises (t20 a0 a1))
(step t22 (cl (= sa$ (arg$ ta$))) :rule resolution :premises (t5 t21))
(step t23 (cl (not (= t2$ (arg$ ta$))) (not (= sa$ (arg$ ta$))) (= t2$ sa$)) :rule eq_transitive)
(step t24 (cl) :rule resolution :premises (t23 a1 t22 a3))
